Before you turn this lab in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
DRIVER = "Aubrey"
NAVIGATOR = "Jonathan"
```

# TensorFlow + Keras Lab

Welcome to the TensorFlow + Keras lab! By the end of this lab you will have

- Optimized and evaluated a keras model with TensorFlow

Let's get started!

# Autograder Variables

The following code defines variables owned by the autograder.

```python
import numpy as np
```

```python
def passed(): print('✅')
```

# Set the Stage

The following code loads data and normalizes it and defines a TensorFlow session and registers it with keras and defines a keras model.

```python
import tensorflow as tf
from keras.datasets import boston_housing
from keras.models import Sequential
from keras import backend as K
from keras.layers import Dense

[X_boston, y_boston], _ = boston_housing.load_data()
X_boston = (X_boston - X_boston.mean(axis=0)) / X_boston.std(axis=0)
y_boston = (y_boston - y_boston.mean()) / y_boston.std()

sess = tf.Session()
K.set_session(sess)

model = Sequential()
model.add(Dense(units=1, input_shape=X_boston[0].shape))
model.summary()
```

Output:
```
_________________________________________________________________
Layer (type)                 Output Shape              Param #   
=================================================================
dense_4 (Dense)              (None, 1)                 14        
=================================================================
Total params: 14.0
Trainable params: 14
Non-trainable params: 0.0
_________________________________________________________________
```

```python
model.output
```

Output:
```
<tf.Tensor 'dense_4/BiasAdd:0' shape=(?, 1) dtype=float32>
```

```python
y_boston = y_boston.reshape(-1,1)
```

```python
X_boston.shape
```

Output:
```
(404, 13)
```

# Task

- Define tensors `labels` and `loss` which hold your labels and aggregated loss for each training example in your minibatch respectively

```python
from keras.objectives import mean_squared_error
```

```python
y_boston[0].shape
```

Output:
```
(1,)
```

<!-- grade_id: tensors_solution, points: 0, grade: false, solution: true -->
```python
labels = tf.placeholder(tf.float32, shape=model.output_shape)
loss = tf.reduce_mean(mean_squared_error(y_boston, model.output))

```

```python
labels
```

Output:
```
<tf.Tensor 'Placeholder_5:0' shape=(?, 1) dtype=float32>
```

# Tensor Tests

```python
type(loss)
```

Output:
```
tensorflow.python.framework.ops.Tensor
```

<!-- grade_id: tensors_tests, points: 1, grade: true, solution: false -->
```python
assert loss.shape.as_list() == []
assert type(loss) == tf.Tensor
assert type(labels) == tf.Tensor
assert labels.op.type == 'Placeholder'
assert loss.dtype == tf.float32
assert loss.shape == []
assert loss.op.type == 'Mean'

passed()
```

Output:
```
✅
```

# Task

- Define a TensorFlow optimizer and initialize all variables

# Requirments

- Define a tensor `train_step` which runs a step of `Adam` to minimize `loss`
- Define a tensor `init` which initializes all `tf.Variable`s and run it

```python

```

<!-- grade_id: optimizer_solution, points: 0, grade: false, solution: true -->
```python
train_step = tf.train.AdamOptimizer().minimize(loss)
init = tf.global_variables_initializer()
sess.run(init)
```

# Optimizer Tests

<!-- grade_id: optimizer_tests, points: 1, grade: true, solution: false -->
```python
assert type(train_step) == tf.Operation
assert type(train_step) == tf.Operation
assert train_step.type == 'NoOp'
assert init.type == 'NoOp'

passed()
```

Output:
```
✅
```

# Task

- Optimize `model` so the loss gets below 1.0 on `(X_boston, y_boston)`

<!-- grade_id: optimize_solution, points: 0, grade: false, solution: true -->
```python
with sess.as_default():
    i=0
    while i < 100:
        l, _ = sess.run(fetches=[loss,train_step], feed_dict={model.input:X_boston,labels:y_boston})
        print(l)
        i+=1
```

Output:
```
0.582464
0.579509
0.576577
0.573667
0.570779
0.567912
0.565067
0.562242
0.559438
0.556655
0.553892
0.551149
0.548425
0.54572
0.543035
0.540368
0.537721
0.535091
0.53248
0.529886
0.52731
0.524752
0.522211
0.519686
0.517179
0.514688
0.512214
0.509756
0.507314
0.504888
0.502477
0.500082
0.497702
0.495338
0.492988
0.490653
0.488333
0.486028
0.483736
0.481459
0.479196
0.476947
0.474711
0.472489
0.470281
0.468085
0.465903
0.463735
0.461579
0.459435
0.457305
0.455187
0.453081
0.450988
0.448907
0.446838
0.444781
0.442736
0.440702
0.438681
0.43667
0.434672
0.432684
0.430708
0.428743
0.426789
0.424846
0.422914
0.420992
0.419081
0.417181
0.415291
0.413412
0.411543
0.409684
0.407836
0.405997
0.404168
0.40235
0.400541
0.398742
0.396952
0.395173
0.393402
0.391642
0.38989
0.388148
0.386415
0.384691
0.382976
0.381271
0.379574
0.377886
0.376207
0.374537
0.372875
0.371222
0.369578
0.367942
0.366315
```

# Evaluation Test

```python
X_boston.shape
```

Output:
```
(404, 13)
```

<!-- grade_id: evaluation_test, points: 1, grade: true, solution: false -->
```python
best_loss = sess.run(fetches=loss, feed_dict={model.input: X_boston, labels: y_boston})
assert best_loss < 1.0

passed()
```

Output:
```
✅
```

```python

```