Before you turn this lab in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
import numpy as np
from IPython.display import display
from keras.datasets import mnist
from keras.preprocessing.image import array_to_img
import numpy as np
```

```python
DRIVER = "Carles"
NAVIGATOR = "Jonathan"
```

# Autoencoder Lab

Welcome to the autoencoder lab! By the end of this lab you will have

- Created and tuned an autoencoder in Keras

*The code in the lab was based on [this blogpost](https://blog.keras.io/building-autoencoders-in-keras.html)*.

# Auto-Grader Variables

```python
ENCODING_DIM = 32

def passed(): print('✅')
```

# Load MNIST

The following code loads MNIST and normalizes it and flattens it.

```python


[X_train, _], [X_test, _] = mnist.load_data()
X_train = np.array([x.flatten() for x in X_train.astype('float32')/255.])
X_test = np.array([x.flatten() for x in X_test.astype('float32')/255.])

for x in X_train[:5]:
    img = array_to_img(x.reshape(28, 28, 1))
    display(img)
```

# Task

- Define an encoder which maps a flattened MNIST image to a vector of size `ENCODING_DIM`

# Requirements

- Use a keras sequential model
- Do not compile your model!

<!-- grade_id: encoder_soution, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

# Encoder Tests

```python
from keras.layers import Dense
from keras.models import Sequential
from keras.regularizers import l1
from IPython.display import SVG
from keras.utils.vis_utils import model_to_dot

```

```python
!pip install graphviz
```

<!-- grade_id: encoder_tests, points: 0, grade: false, solution: true -->
```python
ENCODING_DIM = 32

encoder = Sequential(name='encoder')
encoder.add(Dense(units=ENCODING_DIM, activity_regularizer=l1(1e-5), input_shape=[784]))
#SVG(model_to_dot(encoder).create(prog='dot', format='svg'))
```

# Task

- Define a decoder which maps a compressed version of an image back to a flattened image

# Requirements

- Use a keras sequential model
- Do not compile your model!

<!-- grade_id: decoder_solution, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

# Decoder Tests

<!-- grade_id: decoder_tests, points: 1, grade: true, solution: false -->
```python
import keras

assert type(decoder) == keras.models.Sequential
assert decoder.input_shape == (None, ENCODING_DIM)
assert decoder.output_shape == (None, 784)
assert not hasattr(encoder, 'optimizer')
decoder.predict(encoder.predict(X_train[:1]))

passed()
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-27-dcf2986dbbdd> in <module>()
      1 import keras
      2 
----> 3 assert type(decoder) == keras.models.Sequential
      4 assert decoder.input_shape == (None, ENCODING_DIM)
      5 assert decoder.output_shape == (None, 784)

NameError: name 'decoder' is not defined
```

# Task

- Define an autoencoder by hooking together `encoder` and `decoder` via the keras functional API

# Requirements

- Compile the model with a binary crossentropy loss on a per-pixel basis

<!-- grade_id: autoencoder_solution, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

# Autoencoder Tests

<!-- grade_id: autoencoders_tests, points: 1, grade: true, solution: false -->
```python
import numpy as np

assert type(autoencoder) == keras.engine.training.Model
assert autoencoder.input_shape == encoder.input_shape
assert autoencoder.output_shape == decoder.output_shape
assert hasattr(autoencoder, 'optimizer')
X = X_train[:1]
X_encoded = encoder.predict(X)
assert np.all(decoder.predict(X_encoded) == autoencoder.predict(X))
assert len(autoencoder.layers) == 3
assert type(autoencoder.layers[1]) == keras.models.Sequential
assert type(autoencoder.layers[2]) == keras.models.Sequential
assert type(autoencoder.layers[0]) == keras.engine.topology.InputLayer

passed()
```

# Task

- Optimize your autoencoder on the training data

<!-- grade_id: autograder_solution, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

# Task

- Visualize the performance of your autoencoder

# Requirements

- Convert each image in `X_test` to a `PIL` image and save it into a python list called `test_imgs`
- Run each image in `X_test` through your autoencoder then convert it back into a `PIL` image and save it into a python list called `reconstructed_imgs`

<!-- grade_id: autoencoder_visualization_solution, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

# Autoencoder Visualization Tests

<!-- grade_id: autoencoder_visualization_tests, points: 1, grade: true, solution: false -->
```python
import PIL

assert len(test_imgs) == len(X_test)
for test_img, reconstructed_img in zip(test_imgs, reconstructed_imgs):
    assert type(test_img) == PIL.Image.Image
    assert type(reconstructed_img) == PIL.Image.Image
    assert test_img.size == (28, 28)
    assert reconstructed_img.size == (28, 28)
    
passed()
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-24-2ad5fa2bfccf> in <module>()
      1 import PIL
      2 
----> 3 assert len(test_imgs) == len(X_test)
      4 for test_img, reconstructed_img in zip(test_imgs, reconstructed_imgs):
      5     assert type(test_img) == PIL.Image.Image

NameError: name 'test_imgs' is not defined
```

# Challenge Activities

- Add a sparsity constraint to the encoder
- Make the encoder deeper
- Make the decoder deeper
- Add noise to your data samples

```python
noise_factor = .4

```

```python
X_train_noisy = X_train + noise_factor * np.random.normal(loc=0.0, scale=1.0, size=X_train.shape) 
X_test_noisy = X_test + noise_factor * np.random.normal(loc=0.0, scale=1.0, size=X_test.shape) 
X_train_noisy = np.clip(X_train_noisy, 0., 1.)
X_test_noisy = np.clip(X_test_noisy, 0., 1.)
```

```python
PIL.Image (X_train_noisy[0].reshape(1,28,28))
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-23-279d0022b48d> in <module>()
----> 1 PIL.Image (X_train_noisy[0].reshape(1,28,28))

NameError: name 'PIL' is not defined
```

```python

```